# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
# Complexity 2n+10 implementation
def same_digits(x, y):
    n1 = str(x)
    n2 = str(y)
    count_nums1 = [0] * 10
    count_nums2 = [0] * 10
    is_matching = True
    
    for i in n1:
        count_nums1[int(i)] += 1
    
    for i in n2:
        count_nums2[int(i)] += 1
    
    i = 0
    while i < len(count_nums1) and is_matching:
        if count_nums1[i] != count_nums2[i]:
            is_matching = False
            
        i += 1
    
    return is_matching
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
has_found_answer = False
i = 2

while not has_found_answer:
    if same_digits(i, i * 2) and same_digits(i, i * 3) and same_digits(i, i * 4):
        print(i)
        has_found_answer = True
    
    i += 1
```

Output:
```
142857
```